## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

Some considerations with regard to how these notebooks will be graded:

1. You can add more notebook cells or edit existing notebook cells other than "# YOUR CODE HERE" to test out or debug your code. We actually highly recommend you do so to gain a better understanding of what is happening. However, during grading, **these changes are ignored**. 
2. You must ensure that all your code for the particular task is available in the cells that say "# YOUR CODE HERE"
3. Every cell that says "# YOUR CODE HERE" is followed by a "raise NotImplementedError". You need to remove that line. During grading, if an error occurs then you will not receive points for your work in that section.
4. If your code passes the "assert" statements, then no output will result. If your code fails the "assert" statements, you will get an "AssertionError". Getting an assertion error means you will not receive points for that particular task.
5. If you edit the "assert" statements to make your code pass, they will still fail when they are graded since the "assert" statements will revert to the original. Make sure you don't edit the assert statements.
6. We may sometimes have "hidden" tests for grading. This means that passing the visible "assert" statements is not sufficient. The "assert" statements are there as a guide but you need to make sure you understand what you're required to do and ensure that you are doing it correctly. Passing the visible tests is necessary but not sufficient to get the grade for that cell.
7. When you are asked to define a function, make sure you **don't** use any variables outside of the parameters passed to the function. You can think of the parameters being passed to the function as a hint. Make sure you're using all of those variables.
8. Finally, **make sure you run "Kernel > Restart and Run All"** and pass all the asserts before submitting. If you don't restart the kernel, there may be some code that you ran and deleted that is still being used and that was why your asserts were passing.

# Deep Learning

In this exercise, you will use a deep neural network to predict the values of houses based on some provided input data. You will use keras to build the model. Below is a description of how the keras models are set up.

**Please read more about the data [here](https://www.cs.toronto.edu/~delve/data/boston/bostonDetail.html).**

```python
import tensorflow as tf
import tensorflow.keras as keras
import numpy as np
from tensorflow.keras import Sequential
from tensorflow.keras.layers import Dense, Dropout

np.random.seed(0)
```

```python
(x_train, y_train), (x_test, y_test) = tf.keras.datasets.boston_housing.load_data()
```

```python
x_train.shape
```

Output:
```
(404, 13)
```

```python
y_train.shape
```

Output:
```
(404,)
```

```python
y_train.mean(), y_train.std()
```

Output:
```
(22.395049504950492, 9.199035423364862)
```

The keras model consists of multiple parts:

1. Construct the model layers, neurons per layer, activation function
1. Determine the loss function, metrics and optimization method
1. Fit the model to some data
1. Evaluate the model using the same metric

Some relevant docs:
 - [initializers](https://keras.io/initializers/)
 - [loss functions](https://www.tensorflow.org/api_docs/python/tf/keras/losses)
 - [regularizations](https://keras.io/regularizers/)
 - [optimizers](https://keras.io/optimizers/)
 - [metrics](https://www.tensorflow.org/api_docs/python/tf/keras/metrics)


First, to construct a model, use the [Sequential](https://keras.io/getting-started/sequential-model-guide/) object. You can pass multiple layers to the sequential object. For this exercise, we will only be using the [Dense](https://keras.io/layers/core/#dense) layers.

<!-- grade_id: cell-51879de7032d8823, points: 0, grade: false, solution: true -->
```python
# Create 3 hidden layers with 10 neurons each, and an output layer with 1 neuron. Store the layers
# in a list with the variable name "layers". Pass the list to keras.Sequential and save the
# retuned model to the variable name "model".
#
# Use any activation you like such as "relu" or "tanh", you can alternate for each layer
# For your first run, try using the linear activation and then see if modifying the activations improved the result.
# Optional - give each layer a name and see how that shows up in the summary.
layers = []
layers.append(Dense(10, activation = 'relu', input_shape = (13,) ,name = 'layer1'))
layers.append(Dense(10, activation = 'relu', name = 'layer2'))
layers.append(Dense(10, activation = 'relu', name = 'layer3'))
layers.append(Dense(1, activation = 'relu', name = 'out'))
model = keras.Sequential(layers)

```

<!-- grade_id: cell-ee74eff075f222ed, points: 5, grade: true, solution: false -->
```python
assert isinstance(model, Sequential)
assert len(layers) == 4
for i,layer in enumerate(layers):
    assert isinstance(layers[i], Dense) 
    assert layer.weights[1].shape == [10,10,10,1][i]
```

In TensorFlow, models are "compiled" before training. Compiling specifies the type of optimizer (e.g., gradient descent) and loss function, and creates code that will run efficiently on your hardware during training and model prediction.

Review the models' .compile() methods to create code in the cell below.
https://www.tensorflow.org/api_docs/python/tf/keras/Model#compile


<!-- grade_id: cell-9c5dcca57115d02a, points: 0, grade: false, solution: true -->
```python
# Apply the model's .compile() method to set it up for training.

# Set up the model to do the following:
# - use stochastic gradient descent to fit the model
# - use mean absolute error as its loss function
# - use mean absolute error as one of its metrics
# - use mean squared error as one of its metrics
model.compile(optimizer = 'SGD', loss = 'mae', metrics = ['mae','mse'])
```

<!-- grade_id: cell-0746e29c34118826, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.SGD)
assert model.loss in ["mae", "mean_absolute_error"]
```

```python
# Now fit the model
model.fit(x_train, y_train, epochs=50)
```

Output:
```
Epoch 1/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
Epoch 2/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
Epoch 3/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3951 - mae: 22.3951 - mse: 586.1605
Epoch 4/50
13/13 [==============================] - 0s 2ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
Epoch 5/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1606
Epoch 6/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
Epoch 7/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3951 - mae: 22.3951 - mse: 586.1605
Epoch 8/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
Epoch 9/50
13/13 [==============================] - 0s 1ms/step - loss: 22.3950 - mae: 22.3950 - mse: 586.1605
E
```

```python
model.summary()
```

Output:
```
Model: "sequential"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
layer1 (Dense)               (None, 10)                140       
_________________________________________________________________
layer2 (Dense)               (None, 10)                110       
_________________________________________________________________
layer3 (Dense)               (None, 10)                110       
_________________________________________________________________
out (Dense)                  (None, 1)                 11        
=================================================================
Total params: 371
Trainable params: 371
Non-trainable params: 0
_________________________________________________________________
```

```python
# Here we can evaluate how our model does based on the test data
model.evaluate(x_test, y_test)
```

Output:
```
4/4 [==============================] - 0s 2ms/step - loss: 23.0784 - mae: 23.0784 - mse: 615.8578
[23.078432083129883, 23.078432083129883, 615.8577880859375]
```

Now let's try another optimizer instead of stochastic gradient descent (SGD). [Adam](https://keras.io/optimizers/#adam) is the recommended default for training neural networks since it usually performs quite well. In the next cell, compile the model with Adam instead of SGD and use the same loss and metrics. After compiling, fit the data for as many epochs as you think it takes to see the value start to converge.

<!-- grade_id: cell-bee688f3d919be61, points: 0, grade: false, solution: true -->
```python
# Compile the model using Adam
model.compile(optimizer = 'adam', loss = 'mae', metrics = ['mae','mse'])
```

```python
model.summary()
```

Output:
```
Model: "sequential"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
layer1 (Dense)               (None, 10)                140       
_________________________________________________________________
layer2 (Dense)               (None, 10)                110       
_________________________________________________________________
layer3 (Dense)               (None, 10)                110       
_________________________________________________________________
out (Dense)                  (None, 1)                 11        
=================================================================
Total params: 371
Trainable params: 371
Non-trainable params: 0
_________________________________________________________________
```

<!-- grade_id: cell-fe8cd91b8b780964, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.Adam)
assert model.loss in ["mae", "mean_absolute_error"]
```

```python
model.evaluate(x_test, y_test)
```

Output:
```
4/4 [==============================] - 0s 2ms/step - loss: 23.0784 - mae: 23.0784 - mse: 615.8578
[23.078432083129883, 23.078432083129883, 615.8577880859375]
```

Now recreate the model with dropout layers. Add two dropout layers, one after the first layers of neurons and one after the second layer of neurons.
You may use your existing ```layers``` list and insert the new layers, or create a new list of layers from scratch. Then construct the model as before.
Select a low value of dropout (say, <0.1) that results in a good score.

<!-- grade_id: cell-d5040e44398f8dc9, points: 0, grade: false, solution: true -->
```python
layers.insert(1, Dropout(0.08, name = 'drop1'))
layers.insert(3, Dropout(0.08, name = 'drop2'))
model = keras.Sequential(layers)
```

<!-- grade_id: cell-b6c7e6e2ddc6742a, points: 5, grade: true, solution: false -->
```python
assert len(layers) == 6
assert isinstance(layers[1], Dropout)
assert isinstance(layers[3], Dropout)
for i,layer in enumerate(layers):
    if i not in [1,3]:
        assert isinstance(layers[i], keras.layers.Dense) 
        assert layer.weights[1].shape == [10,0,10,0,10,1][i]
```

```python
model.compile(optimizer='adam', loss='mae', metrics=['mae', "mse"])
model.fit(x_train, y_train, epochs=500, verbose=0)
model.summary()
```

Output:
```
Model: "sequential_1"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
layer1 (Dense)               (None, 10)                140       
_________________________________________________________________
drop1 (Dropout)              (None, 10)                0         
_________________________________________________________________
layer2 (Dense)               (None, 10)                110       
_________________________________________________________________
drop2 (Dropout)              (None, 10)                0         
_________________________________________________________________
layer3 (Dense)               (None, 10)                110       
_________________________________________________________________
out (Dense)                  (None, 1)                 11        
======================================================
```

```python
model.evaluate(x_test, y_test)
model.evaluate(x_test, y_test)
```

Output:
```
4/4 [==============================] - 0s 2ms/step - loss: 4.6383 - mae: 4.6383 - mse: 40.8839
4/4 [==============================] - 0s 2ms/step - loss: 4.6383 - mae: 4.6383 - mse: 40.8839
[4.638320446014404, 4.638320446014404, 40.88386154174805]
```

Select a dropout rate that gets an okay score.

<!-- grade_id: cell-538bede4947fe364, points: 0, grade: true, solution: false -->
```python
assert model.evaluate(x_test, y_test)[0] < 7
```

Output:
```
4/4 [==============================] - 0s 2ms/step - loss: 4.6383 - mae: 4.6383 - mse: 40.8839
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    return "All good!"  
```

<!-- grade_id: feedback-tests, points: 0, grade: true, solution: false -->
```python
print(feedback())
```

Output:
```
All good!
```